```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment4.ipynb")
```

# Assignment 4: Modeling and Optimization

Mathematical modeling of a problem at hand give us a systematic way of finding a solution. For example, a maximum likelihood estimator (assuming it exists), $\hat\theta$, is a method for finding the parameter that maximizes the likelihood of the data $L_n$:
$$L_{n}(\hat{\theta} ; x_1, x_2, \dots, x_n)=\max _{\theta \in \Theta} L_{n}(\theta ; x_1, x_2, \dots, x_n)$$
Data $x_1, x_2, \dots$, set of feasible parameters $\Theta$, and likelihood function $L_n$ are given. We find the parameter that "best" describe the data in the context of the likelihood function.

Many other applications of optimization exists, and this assignment will give a hands-on introduction to a simple linear programming problem.

```python
import cvxpy as cp
import numpy as np
import pandas as pd
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
sns.set_style("whitegrid")
```

<!-- BEGIN QUESTION -->


## Questions 1-3: Resource Allocation Problem

You are in charge of a company that makes two hot sauces: $x_1$ liters of Kapatio and $x_2$ liters of Zriracha. We will use optimization technique to find the "best" manufacturing strategy given our resource constraints.

First, we need to define what we mean by "best" strategy. In this scenario, the goal is to obtain the highest revenue possible. While doing so, there are resource constraints we must satisfy. 

For example, to manufacture the two hot sauce products, different amounts of pepper and vinegar are needed, and each resource is limited.

![ingridient table](table.png)

## Question 1: Resource Constraints

### Question 1.a: Modeling Resource Usage

What is the equation for the amount of pepper needed to manufacture $x_1$ and $x_2$. What is the equation for the amount of vinegar?
(Use [Mathpix](https://mathpix.com/) to write equations, looking out for possible pdf export failures with latex syntax errors.)


**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

### Question 1.b:  Resource Usage vs. Total Resource Constraint

Total amount of pepper needed cannot exceed total available. Write down the inequality expressing this relationship. Do the same for vinegar. These inequalities are your resource constraints.
Additinally, variables $x_1$ and $x_2$ are non-negative: i.e. amount of manufactured goods cannot be negative.

Rewrite the system of constraint inequalities into a matrix inequality: $Ax\leq b$, where $x=(x_1, x_2)^T$. Arrange rows of $A$ and $b$ such that:

* Row 1: total pepper amount constraint
* Row 2: total vinegar amount constraint
* Row 3: Kapatio non-negativity constraint
* Row 4: Zriracha non-negativity constraint

Less than symbol in $Ax\leq b$ means element-wise.

</div>

**SOLUTION**

<!-- END QUESTION -->

<div class="alert alert-block alert-success">
Define matrix `A1` and vector `b1` according to matrix inequality above.
</div>

```python
A1 = ...
b1 = ...
```

```python
grader.check("q1b2")
```

### Visualizing Feasible Region <a class="tocSkip">
In a 2-dimensional plot, we will visualize the area that satisfies both of the resource constraints. Draw $x_1$ on the horizontal axis and $x_2$ on the vertical axis. 

There will be two main components to the plot:
* **Lines** indicating constraint boundaries:  
    e.g. the constraint $x_2\geq 0$ has boundary at $x_2 = 0$.
* **Shaded area** indicating feasible regions:  
    e.g., the whole region $x_2 > 0$ is to be shaded _if_ $x_2\geq 0$ was the only constraint. We will use shading to indicate the region where _all_ constraints are satisfied.

```python
x1_line = np.linspace(-1, 10, 500)
x2_line = np.linspace(-1, 10, 500)
```

<div class="alert alert-block alert-success">

### Question 1.c: Feasible Region Boundary

In a list named `boundary`, create four data frames for each equality in $Ax=b$. These lines indicate where the feasible area ends. Set column names as

* `$x_1$`
* `$x_2$`
* `constraints`

Note the use of latex codes (feel free to use [Mathpix](https://mathpix.com/) but be careful with pdf compilation errors). 

#### Toy Example: Drawing Boundaries <a class="tocSkip">

Here is a **toy example** of drawing two constraint boundaries by constructing data frames:

</div>

```python
z1_line = np.linspace(-1, 10, 500)
z2_line = np.linspace(-1, 10, 500)

boundary = [
    pd.DataFrame({
        '$z_1$': np.ones_like(z2_line)*1,
        '$z_2$': z2_line,
        'constraint': '$z_1\\geq 1$'
    }),
    pd.DataFrame({
        '$z_1$': z1_line,
        '$z_2$': np.ones_like(z1_line)*2,
        'constraint': '$z_2\\geq 2$'
    }),
]
fig, ax = plt.subplots(figsize=(6, 6))
sns.lineplot(x='$z_1$', y='$z_2$', hue='constraint', data=pd.concat(boundary).reset_index(), ax=ax).axvline(1)
plt.xlim(-1, 6)
plt.ylim(-1, 6)
plt.show()
```

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">
Sometimes, things just do not work as expected. 

In the toy example code, 
```
sns.lineplot(x='$z_1$', y='$z_2$', hue='constraint', data=pd.concat(boundary), ax=ax).axvline(1)
```
what seems strange about the plotting command? Why was the strange code necessary?
</div>

**SOLUTION**

#### Example: Resource Constraint Boundary  <a class="tocSkip">

Now, create a data frame for the non-negativity constraint $x_2\geq 0$ as follows:

```python
pd.DataFrame({'$x_1$':x1_line,               ## x_1 can take on any value
              '$x_2$':x1_line * 0.0,         ## x_2 = 0
              'constraint':'$x_2 \\geq 0$'}), ## constraint equation for labeling
```

<!-- END QUESTION -->

<div class="alert alert-block alert-success">
Create a list named `boundary` containing four data frames (each corresponding to a constraint). Concatenate data frames in `boundary` to one data frame named `hull`.
</div>

```python
boundary = [
    pd.DataFrame({'$x_1$': ...,
                  '$x_2$': ...,
                  'constraint':'$... x_1 + ... x_2 \\leq ...$'}),
    pd.DataFrame({'$x_1$': ...,
                  '$x_2$': ...,
                  'constraint':'$... x_1 + ... x_2 \\leq ...$'}),
    pd.DataFrame({'$x_1$': ..., 
                  '$x_2$': ...,
                  'constraint':'$... x_1\\geq ...$'}),
    pd.DataFrame({'$x_1$':x1_line,
                  '$x_2$':x1_line * 0.0,
                  'constraint':'$x_2\\geq 0$'})
]
hull = ...
```

```python
grader.check("q1c2")
```

<div class="alert alert-block alert-success">

### Question 1.d: Interior of Feasible Region

Previous question prepared constraint boundaries, $Ax = b$. In this question, we calculate the interior of the feasible region, which will be shaded in the visualization. First, create a 2-d array of $x_1$ and $x_2$ values. If a point $(x_1, x_2)$ satisfies _every_ constraint, the point will be colored grey.

For example, in order to shade $\{x_1: x_1\geq 1\}\cap\{x_2: x_2\geq 2\}$, we can use the `imshow` method. 

</div>

```python
z1_line = np.linspace(-1, 6, 10)
z2_line = np.linspace(-1, 6, 10)
z1_grid, z2_grid = np.meshgrid(z1_line, z2_line)

fig, az = plt.subplots(figsize=(6, 6))
az.imshow(
    ((z1_grid >= 1) & (z2_grid >= 2)).astype(int),
    origin='lower',
    extent=(z1_grid.min(), z1_grid.max(), z2_grid.min(), z2_grid.max()),
    cmap="Greys", alpha=0.3, aspect='equal' 
)
plt.xlim(-1, 6)
plt.ylim(-1, 6)
plt.show()
```

<!-- BEGIN QUESTION -->

By dissecting the command below and reading the documentation, report what each of the following lines does:

* `((y1_grid >= 1) & (y2_grid >= 2)).astype(int)` (What is the output of running this command?)  
* `origin='lower'`  
* `extent=(y1_grid.min(), y1_grid.max(), y2_grid.min(), y2_grid.max())`  
* `cmap='Greys'`  
* `alpha=0.3`  
* `aspect='equal'`  

**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

### Question 1.e: Visualizing the Feasible Region

Finally, create a figure that shows constraint boundaries and the interior region shaded with a light grey color.

Your output will look like this:  
![hull](feasible.png)

</div>

```python
fig, ax = plt.subplots(figsize=(6, 6))
x1_grid, x2_grid = np.meshgrid(x1_line, x2_line)

ax.imshow(
    (
        # BEGIN SOLUTION
        (A1[0,0]*x1_grid + A1[0,1]*x2_grid <= b1[0]) & # Pepper constraints
        (A1[1,0]*x1_grid + A1[1,1]*x2_grid <= b1[1]) & # Vinegar constraints
        (A1[2,0]*x1_grid + A1[2,1]*x2_grid <= b1[2]) & 
        (A1[3,0]*x1_grid + A1[3,1]*x2_grid <= b1[3])   # non-negativity constraints 
        # END SOLUTION
    ),
    origin='lower',
    extent=(x1_grid.min(), x1_grid.max(), x2_grid.min(), x2_grid.max()),
    cmap="Greys", alpha = 0.3, aspect='equal')


# ax = sns.lineplot(???).axvline(???)
sns.lineplot(x='$x_1$', y='$x_2$', hue='constraint', data=hull.reset_index(), ax=ax).axvline(0, color='green') # SOLUTION

plt.xlim(-1, 6)
plt.ylim(-1, 6)
#fig.savefig("files/feasible.png") # SOLUTION NO PROMPT
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

In the context of linear programming, $Ax\leq b$ is called the _feasible region_ (including the appropriate sections of the boundaries). Denote the (shaded) feasible region as set $C$. Points $(x_1,x_2)\in C$ satisfy all of the constraints.

Describe in plain words the feasible region in the context of hot sauce manufacturing. Specifically, which constraint is violated (if any) by a point at:

* $(x_1,x_2) = (4, 1)$
* $(x_1,x_2) = (0, 5)$  
* $(x_1,x_2) = (3, 4)$  
</div>

**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

## Question 2: Objective Function

### Question 2.a: Defining Objective Function

Suppose the hot sauces are sold at the same price: \\$5 per liter. 

What is the equation $f(x)$ for the total revenue as a function of $x_1$ and $x_2$? 

The function $f(x)$ is called the objective function.

</div>

**SOLUTION**

<!-- END QUESTION -->

<div class="alert alert-block alert-success">
Objective function $f(x)$ is a linear function in $x$. Therefore, $f(x)$ is a 2-dimesional hyperplane. Note that each value of $f(x)$ defines a line in $(x_1, x_2)$ plane.

For example, $f(x)= 0 = c_1x_1 + c_2x_2$ defines a line. A subspace of equal function value is sometimes referred to as a _level set_ or a _contour line_ when visualized.

First, create a numpy array of prices `c` for the two hot sauces, $x_1$ and $x_2$. Then, create a list `f_vals` containing four data frames of contour lines, $f(x) \in \{0, 10, 20, 30\}$. by creating one data frame for each contour line. 
</div>

```python
c = ...

fig, ax = plt.subplots(figsize=(6, 6))
contours = [
    pd.DataFrame({
        '$x_1$': x1_line,
        '$x_2$': ( 0 - c[0]*x1_line)/c[1],
        '$f(x)$': 0
    }),
    ...
]
f_vals = pd.concat(contours)

# ax = sns.lineplot(???)
ax = ...

plt.xlim(-1, 6)
plt.ylim(-1, 6)
plt.show()
```

```python
grader.check("q2a2")
```

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

### Question 2.b: Direction of Steepest Increase

Since we want to maximize revenue, we want to increase our objective function as much as possible. Analogous to the minimization example given in a previous lecture, we can repeatedly move in the direction of function increase. In order to determine such direction, compute the gradient of $f(x)$ at $x=(0,0)^T$:
$$\nabla_x f(x) = \begin{pmatrix}\frac{\partial f(x)}{\partial x_1}\\\frac{\partial f(x)}{\partial x_2}\end{pmatrix} $$

</div>

**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

## Question 3: Putting Pieces Together

### Question 3.a: Standard Form of a Linear Programming Problem

Write down the so-called the _standard form_ of a linear programming problem:
$$\begin{aligned}
\max_x & f(x)\\
\text{subject to } & Ax\leq b
\end{aligned}$$

Specifically, write the obejective as an inner product of two vectors: $f(x) = c^T x$, and write the constraint as a vector inequality involving a matrix-vector prduct: $Ax\leq b$, where $A$ is a 4-by-2 matrix.

</div>

**SOLUTION**

<!-- END QUESTION -->

<div class="alert alert-block alert-success">

### Question 3.b: Computing the Numerical Solution

Therefore, _maximizing_ the revenue is a search over the feasible region for the best point $x^*=(x_1^*, x_2^*)$ that gives the largest revenue. On the otherhand, any _infeasible_ point _not_ in the feasible region cannot be a solution to the constrained optimization problem.

Notationally, the following expression means the same thing:
$$x^* = \arg\max_{\{x: Ax\leq b\}} f(x)$$

Using [CVXPY](https://www.cvxpy.org/), solve for the resource allocation problem with constraints.

</div>

```python
# define variables
x = ...

# define the linear program
problem = cp.Problem(
    ...
)

fstar1 = problem.solve()                                                 # maximum attained function value
xstar1 = pd.DataFrame(x.value.reshape(1, 2), columns=['$x_1$', '$x_2$']) # maximizer x for f
```

```python
grader.check("q3b")
```

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

### Question 3.c: Plotting the optimal solution

</div>

```python
fig, ax = plt.subplots(figsize=(6, 6))
x1_grid, x2_grid = np.meshgrid(x1_line, x2_line)
ax.imshow(
    (
        (A1[0,0]*x1_grid + A1[0,1]*x2_grid <= b1[0]) & # Pepper constraints
        (A1[1,0]*x1_grid + A1[1,1]*x2_grid <= b1[1]) & # Vinegar constraints
        (A1[2,0]*x1_grid + A1[2,1]*x2_grid <= b1[2]) & 
        (A1[3,0]*x1_grid + A1[3,1]*x2_grid <= b1[3])   # non-negativity constraints 
    ),
    origin='lower',
    extent=(x1_grid.min(), x1_grid.max(), x2_grid.min(), x2_grid.max()),
    cmap="Greys", alpha = 0.3, aspect='equal' 
)
sns.scatterplot(x='$x_1$', y='$x_2$', data=xstar1, ax=ax, s=100)
sns.lineplot(x='$x_1$', y='$x_2$', hue='$f(x)$', data=f_vals.reset_index(), ax=ax)
plt.xlim(-1, 6)
plt.ylim(-1, 6)
plt.show()
```

<!-- END QUESTION -->

<div class="alert alert-block alert-success">

## Question 4: Nutrition Problem

During the second world war, the US Army set out to save money without damaging the nutritional health of members of the armed forces. 

> According to [this source](https://ibmdecisionoptimization.github.io/docplex-doc/mp/diet.html), the following problem is a _simple variation of the well-known diet problem that was posed by George Stigler and George Dantzig: how to choose foods that satisfy nutritional requirements while minimizing costs or maximizing satiety._
> 
> _Stigler solved his model "by hand" because technology at the time did not yet support more sophisticated methods. However, in 1947, Jack Laderman, of the US National Bureau of Standards, applied the simplex method (an algorithm that was recently proposed by George Dantzig) to Stigler’s model. Laderman and his team of nine linear programmers, working on desk calculators, showed that Stigler’s heuristic approximation was very close to optimal (only 24 cents per year over the optimum found by the simplex method) and thus demonstrated the practicality of the simplex method on large-scale, real-world problems._
> 
> _The problem that is solved in this example is to minimize the cost of a diet that satisfies certain nutritional constraints._

The file `foods.csv` contains calorie, nutritional content, serving size, and price per serving information about 64 foods. Read it into a data frame named `foods`.
</div>

```python
foods = pd.read_csv('files/foods.csv') # SOLUTION
print(foods)
```

<div class="alert alert-block alert-success">
The file `nutritional_constraints.csv` contains healthy nutritional range constraints. Minimum and maximum allowed nutritional contents can be found in this file. Name the variable `requirements`.
</div>

```python
requirements = pd.read_csv('files/nutritional_constraints.csv') # SOLUTION
print(requirements)
```

<div class="alert alert-block alert-success">
Extract the nutritional content of foods into a 2-d array named `ncontent`. The dimension of `ncontent` DataFrame should be 11 rows and 64 columns.
</div>

```python
ncontent = foods.iloc[:,range(1,12)].T # SOLUTION
print(ncontent)
```

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

### Question 4.a: Define Constraints

To avoid eating the same foods, limit each food intake to be 2 or less. Also, one cannot consume less than zero servings. Furthermore, apply the nutritional constraints as specified in `nutritional_constraints.csv` (assume that the units are the same as food nutritional contents)

Note that a range constraints, e.g., $2000 \leq \text{total calories} \leq 2250$, can be written as two constraints: $\text{total calories} \leq 2250$ and $-\text{total calories} \leq -2000$. Hence, we can rewrite caloric intake constraints as
$$\begin{aligned}
-(\text{calories in frozen broccoli})x_0 - (\text{calories in raw carrots})x_1 - \dots - (\text{calories in bean bacon soup, w/watr})x_{63} = -c^T x&\leq -2000\\
  \text{calories in frozen broccoli})x_0 + (\text{calories in raw carrots})x_1 + \dots + (\text{calories in bean bacon soup, w/watr})x_{63} = c^T x&\leq 2250\\
\end{aligned}$$
where vector $c$ contains calorie information for all 64 foods and $x$ containts servings consumed of each food. Matrix $U$ and vector $w$ would be such that
$$\begin{aligned}
U=
\begin{pmatrix}
-c^T\\
c^T
\end{pmatrix}
\text{ and }
w = \begin{pmatrix}
-2000\\
2250
\end{pmatrix}
\end{aligned},$$
and the matrix-vector inequality would be $Ux\leq w$. Range constraints of each food can be implemented similarly with identity matrices.

Denote nutritional content information from `foods` data frame as $A$ and denote the `Min` and `Max` columns of `requirements` as vector $b_L$ and $b_U$, respectively. Construct $M$ and $d$ in $M x \leq d$ using $I$ (identity matrix), $A$, $b_L$, $b_U$, and other constants, so that all the range constraints are expressed in $M x\leq d$. (This is a theory question. No coding is involved)

</div>

**SOLUTION**

<!-- END QUESTION -->

<div class="alert alert-block alert-success">

### Question 4.b: Create Python Variables

Denote the servings of each food as $x_i$ where $i$ is the row index of each food in `foods` data frame: i.e. $x_0$ indicates number of servings of frozen broccoli, $x_1$ indicates that of raw carrots, etc.

* Create cost vector `cost` that gives per serving cost. 
* Create matrix `M` and `d` that lists nutritional content in the following order:
    * Non-negativity constraint of food consumed: i.e. 0 servings or more
    * Upper limit on food consumed: i.e. 2 servings or less
    * Lower limit on consumption of each nutrition: i.e. following `Min` column
    * Upper limit on consumption of each nutrition: i.e. following `Max` column

</div>

```python
M = np.concatenate([
    ...
])

d = np.concatenate([
    ...
])

cost = ...
```

```python
grader.check("q4b")
```

<div class="alert alert-block alert-success">

### Question 4.c: Solve the Problem

Create cvxpy variable `servings` to represent the number of servings of food, and use cvxpy to solve for the optimal solution. 

</div>

```python
servings = ...

# define the linear program
nutrition_problem = cp.Problem(
    ...
)

# solve linear programming problem with ECOS solver
# https://www.cvxpy.org/tutorial/advanced/index.html?highlight=osqp#choosing-a-solver
fstar2 = ...
xstar2 = servings.value
```

```python
grader.check("q4c")
```

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

### Question 4.d: Interpreting the Results

State the results in the context of the problem. How much of each food was consumed? List the foods and their calculated amounts. What is the total cost of feeding one soldier? 

</div>

**SOLUTION**

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

## Question 5: Assignment Problem (Simulated Case Study)

A startup has $m=10$ engineers and $m=10$ project tasks that must be assigned for next week. Each engineer must be assigned to exactly one task, and each task must be assigned to exactly one engineer.

Let $\mathcal{S}$ be the set of engineers, $\mathcal{D}$ the set of tasks, and let $c_{ij}$ be the cost (time + coordination overhead) of assigning engineer $i\in\mathcal{S}$ to task $j\in\mathcal{D}$. Define

$$
x_{i j}= \begin{cases}1 & \text { if engineer } i \text { is assigned to task } j, \\ 0 & \text { otherwise. }\end{cases}
$$

Then the objective is

$$
\operatorname{minimize} \sum_{i \in \mathcal{S}} \sum_{j \in \mathcal{D}} c_{i j} x_{i j}.
$$

### Question 5.a: Write the equality constraints

Write the equality constraints that enforce:
1. Each engineer is assigned exactly one task.
2. Each task is assigned to exactly one engineer.

</div>

**SOLUTION**

<!-- END QUESTION -->

<div class="alert alert-block alert-success">

The following constructs a simulated assignment cost matrix.

Requirements:
- Set `m = 10`: 10 jobs and tasks.
- Create labels `engineers` and `tasks`
- Use the provided fixed matrix `C`
- Build a DataFrame `cost_df` with engineers as rows and tasks as columns

</div>

```python
m = 10

engineers = [f"E{i+1}" for i in range(m)]
tasks = [f"T{j+1}" for j in range(m)]

C = np.array([[48, 60, 65, 25, 32, 15, 45, 61, 24, 29],
              [37, 65, 54, 64, 13, 28, 11, 33, 72, 64],
              [32, 39, 50, 42, 45, 24, 54, 43, 36, 19],
              [34, 41, 13, 78, 54, 29, 17, 66, 34, 38],
              [60, 58, 63, 67, 80, 38, 47, 44, 23, 62],
              [52, 73, 77, 23, 53, 63, 69, 65, 26, 16],
              [73, 51, 27, 21, 44, 20, 19, 51, 59, 20],
              [16, 36, 52, 28, 68, 59, 32, 27, 46, 51],
              [80, 47, 46, 46, 45, 21, 54, 72, 71, 15],
              [24, 25, 62, 72, 29, 10, 64, 58, 77, 11]])

cost_df = pd.DataFrame(C, index=engineers, columns=tasks)
cost_df
```

<div class="alert alert-block alert-success">

### Question 5.b: Solve the assignment model with CVXPY

Use CVXPY to solve the assignment problem with binary decision matrix `X` of shape `(m, m)`, where `X[i, j] = 1` means engineer `i` is assigned task `j`.

Steps:
1. Create `X = cp.Variable((m, m), boolean=True)`
2. Add row-sum constraints so each engineer gets exactly one task.
3. Add column-sum constraints so each task is assigned exactly once.
4. Minimize total assignment cost `cp.sum(cp.multiply(C, X))`
5. Solve the model and create `assignment_df` as an integer assignment matrix.

</div>

```python
X = ...


constraints = [
    ...,  # each engineer gets one task
    ...   # each task gets one engineer
]
...

assignment_problem = cp.Problem(
    ...
    constraints,
 )

assignment_cost = ...

assignment_df = ...
assignment_df
```

```python
grader.check("q5b")
```

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

### Question 5.c: Interpret the optimal assignment

Using `assigned_pairs` and `assignment_cost`, report:
1. Which engineer is assigned to each task.
2. The total minimized assignment cost.
3. A short interpretation of why this assignment is better than a random assignment.

</div>

**SOLUTION**

<!-- END QUESTION -->

_Blank Cell_

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(filtering=False, run_tests=True)
```

Tests for q1b2 (points: 4)
```
>>> A1.shape
(4, 2)
```
```
>>> b1.shape
(4,)
```
```
>>> A1
array([[ 5,  7],
       [ 4,  2],
       [-1,  0],
       [ 0, -1]])
```
```
>>> b1
array([30, 12,  0,  0])
```


Tests for q1c2 (points: 6)
```
>>> len(boundary)
4
```
```
>>> np.all([type(d) is pd.DataFrame for d in boundary])
np.True_
```
```
>>> type(hull) is pd.DataFrame
True
```
```
>>> np.all([f in hull.columns.values for f in ['$x_1$', '$x_2$', 'constraint']])
np.True_
```


Tests for q2a2 (points: 4)
```
>>> len(contours)
4
```
```
>>> f_vals.shape
(2000, 3)
```
```
>>> (ax.xaxis.get_label_text(), ax.yaxis.get_label_text())
('$x_1$', '$x_2$')
```
```
>>> [f.get_text() for f in ax.legend_.get_texts()]
['0', '10', '20', '30']
```


Tests for q3b (points: 4)
```
>>> np.round(fstar1, 3)
np.float64(23.333)
```
```
>>> np.round((A1 @ x.value).sum(), 3)
np.float64(37.333)
```
```
>>> np.round(x.value, 3)
array([1.333, 3.333])
```


Tests for q4b (points: 4)
```
>>> M.shape
(150, 64)
```
```
>>> d.shape
(150,)
```
```
>>> M.sum()
np.float64(-9.094947017729282e-12)
```
```
>>> d.sum()
np.float64(69338.0)
```
```
>>> cost
array([0.16, 0.07, 0.04, 0.18, 0.02, 0.53, 0.06, 0.31, 0.84, 0.78, 0.27,
       0.24, 0.15, 0.32, 0.49, 0.15, 0.16, 0.05, 0.06, 0.09, 0.16, 0.03,
       0.05, 0.25, 0.16, 0.23, 0.13, 0.08, 0.11, 0.15, 0.27, 0.33, 0.15,
       0.31, 0.28, 0.28, 0.34, 0.32, 0.38, 0.82, 0.52, 0.44, 0.59, 0.83,
       0.31, 0.39, 0.08, 0.17, 0.07, 0.81, 0.45, 0.69, 0.04, 0.22, 0.12,
       0.19, 0.39, 0.67, 0.71, 0.75, 0.39, 0.99, 0.65, 0.67])
```


Tests for q4c (points: 4)
```
>>> np.round(xstar2, 2)
array([0.08, 0.22, 0.  , 0.  , 0.  , 0.  , 2.  , 0.  , 0.  , 0.  , 0.  ,
       0.  , 0.  , 0.  , 0.  , 0.  , 0.  , 2.  , 1.42, 0.  , 0.  , 2.  ,
       0.  , 0.  , 1.93, 0.  , 0.15, 0.25, 0.  , 0.  , 0.  , 0.  , 0.  ,
       0.  , 0.  , 0.  , 0.  , 0.  , 0.  , 0.  , 0.24, 0.  , 0.  , 0.  ,
       0.  , 0.  , 2.  , 0.  , 2.  , 0.  , 0.  , 0.  , 2.  , 0.  , 0.  ,
       0.  , 0.  , 0.  , 0.  , 0.  , 0.  , 0.  , 0.  , 0.  ])
```
```
>>> ncontent @ xstar2 <= requirements['Max'].values
Calories         True
Cholesterol      True
Total_Fat        True
Sodium           True
Carbohydrates    True
Dietary_Fiber    True
Protein          True
Vit_A            True
Vit_C            True
Calcium          True
Iron             True
dtype: bool
```
```
>>> ncontent.shape
(11, 64)
```
```
>>> ncontent @ xstar2 >= requirements['Min'].values
Calories         True
Cholesterol      True
Total_Fat        True
Sodium           True
Carbohydrates    True
Dietary_Fiber    True
Protein          True
Vit_A            True
Vit_C            True
Calcium          True
Iron             True
dtype: bool
```
```
>>> np.round(ncontent @ xstar2, 2)
Calories         2000.00
Cholesterol       127.84
Total_Fat          65.00
Sodium           1216.43
Carbohydrates     300.00
Dietary_Fiber      25.00
Protein            66.51
Vit_A            5000.00
Vit_C              50.00
Calcium           800.00
Iron               29.65
dtype: float64
```


Tests for q5b (points: 6)
```
>>> (assignment_df.shape, np.all(assignment_df.sum(axis=1) == 1), np.all(assignment_df.sum(axis=0) == 1))
((10, 10), np.True_, np.True_)
```
```
>>> np.round(assignment_cost, 3)
np.float64(205.0)
```
```
>>> assigned_pairs = [(engineers[i], tasks[j], C[i, j]) for i in range(m) for j in range(m) if assignment_df.iloc[i, j] == 1]
>>> print(assigned_pairs)
[('E1', 'T6', np.int64(15)), ('E2', 'T5', np.int64(13)), ('E3', 'T8', np.int64(43)), ('E4', 'T3', np.int64(13)), ('E5', 'T9', np.int64(23)), ('E6', 'T4', np.int64(23)), ('E7', 'T7', np.int64(19)), ('E8', 'T1', np.int64(16)), ('E9', 'T10', np.int64(15)), ('E10', 'T2', np.int64(25))]
```
