```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment0.ipynb")
```

<div class="alert alert-block alert-info">

**Instructions**: Double-click this cell to edit it. Replace the text below with your actual information.

### Name: 
### Perm Number: 

</div>

<div class="alert alert-block alert-success">

# Working on PSTAT 134 Assignments 

This notebook will take you through how assignments are structured in this class and help you practice submitting your work on Gradescope.

## Tools in our Workflow

- [**JupyterHub**](https://pstat134.lsit.ucsb.edu/): A centrally managed computational platform where everyone has exactly the same Python environment.
- [**GitHub**](https://github.com/UCSB-PSTAT-134/Winter2026): A version-controlled code repository that tracks all changes to class materials.
- [**Gradescope**](https://ucsb.instructure.com/courses/31641/external_tools/158): The platform where you will submit and autograde your assignments. **Note**: Access Gradescope via the Canvas menu to ensure your accounts are properly linked.
- [**Otter Grader**](https://otter-grader.readthedocs.io/): A framework for creating assignments with embedded tests to evaluate your code.

---

## Preliminaries

### Jupyter Lab Basics

Jupyter notebooks consist of **cells**. There are two main types:
1. **Markdown cells**: For text, links, and math (like this one).
2. **Code cells**: For executing Python code.

- [JupyterLab Documentation](https://jupyterlab.readthedocs.io/en/stable/user/notebook.html)
- [Navigate using the Table of Contents](https://jupyterlab.readthedocs.io/en/stable/user/toc.html)

### Math Notation

You can use LaTeX to write math in Markdown cells.
- Inline math: Wrap in single dollar signs, like $E = mc^2$.
- Block math: Wrap in double dollar signs:
$$\frac{1}{n} \sum_{i=1}^n X_i$$

### How we sync files: NBGitPuller

We use **NBGitPuller** to manage the distribution of notebooks. When you click an assignment link from the course website, NBGitPuller automatically:
1. "Pulls" the latest version of the assignment from our GitHub repository.
2. Merges it with any existing files in your JupyterHub workspace.
3. Opens the notebook for you to work on.

[Learn more about NBGitPuller](https://jupyterhub.github.io/nbgitpuller/index.html)

</div>

<div class="alert alert-block alert-warning">

## How to Test Your Work with Otter Grader

In most assignments, you will see special cells with code, `grader.check(...)`. These lines run embedded tests to to verify your solution against the expected output.

- **Running Tests**: After writing your code in a solution cell, **run the cell** (Shift+Enter). Then, run the test cells that follow it.
- **Feedback**: 
    - **Visible Tests**: If your code is correct, you will see `Test Passed`. If it's wrong, you'll see a descriptive error message.
    - **Hidden Tests**: Some tests are "hidden" and only run once you submit to Gradescope. You won't see their results locally.
- **Check All**: You can run `otter.Notebook().check_all()` if available, but usually, running individual test cells as you go is the best practice.

</div>

<div class="alert alert-block alert-success">

## Problem 1

Write a function `square_num` that calculates the square of a number.

### Explanation

In this question, there are two tests:

There are two tests,

1. One visible test: `square_num(2)`. Visible embedded tests will tell you what code is being executed to evaluate correctness. Additionally, what result is expected vs. what is actually obtained by executing your code.
1. One hidden test: `square_num(12)`. Hidden test will be executed on Gradescope platform, and you will not see any result from this test in your notebook.

### Your Turn

**Test what happens when your function is correct vs. wrong.**

</div>

```python
def square_num(x):
    
    ...

    return ...

```

```python
grader.check("q1")
```

<div class="alert alert-block alert-success">

## Problem 2

Sort the rows of `fruit_info` below by
1. the entries of `color` column in alphabetical order. Save it to a variable named `fruit_by_color`
1. the entries of `fruit` column in alphabetical order. Save it to a variable named `fruit_by_name`

What is the name of the fruit with color pink? Save the name to variable named `pink_fruit` as a string

### Explanation

This is another autograded question.

There are three test cases:
1. Visible test: `fruit_by_color.index.to_list()`
1. Hidden test: `fruit_by_name.index.to_list()`
2. Hidden test: `assert pink_fruit == 'raspberry'`

### Your turn



</div>

```python
import pandas as pd

fruit_info = pd.DataFrame(
    data={'fruit': ['apple', 'orange', 'banana', 'raspberry'],
          'color': ['red', 'orange', 'yellow', 'pink']
          })
fruit_by_color = ...
fruit_by_name = ...
pink_fruit = ...
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-success">

## Problem 3: Simulation and Visualization

Complete the code below to simulate 1000 observations of $x$ and $y$ with a specific dependency structure.

- Variable `x` should contain 1000 iid observations from a **Gamma distribution** with shape parameter 2.
- Variable `y` should be a linear combination of `x` and random noise: $y = -0.5 \cdot x + \epsilon$, where $\epsilon$ follows a **Standard Normal distribution**.

### Why this is different
This is a **manually graded question**. Unlike Problems 1 and 2:
1. You won't see "Test Passed" feedback in the notebook.
2. An instructor or TA will review your code and the resulting plot on Gradescope.
3. The export process will generate a PDF version of your notebook (including this plot) for manual review.

### Your turn
Follow the prompts in the code cell below to generate the joint hexagonal bin plot.

</div>

```python
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
import numpy as np

sns.set_theme(style="ticks")

rs = np.random.RandomState(11)

x = rs.gamma(2, size=...)
y = -.5 * x + rs.normal(size=...)

sns.jointplot(x=x, y=y, kind="hex", color="#4C6DB3")
plt.show()
```

<!-- END QUESTION -->

<div class="alert alert-block alert-info">

## Submission Instructions

Once you have completed all problems in this notebook:

1. **Verify your code**: Restart the kernel (`Kernel` -> `Restart Kernel & Run All Cells`) to ensure everything runs from scratch without errors.
2. **Verify your PDF**: If any PDF is generated, open the generated PDF file and ensure your manually graded questions show up fully and properly.
3. **Submit to Gradescope**: 
    - Navigate to our course on Gradescope through Canvas (or login to Gradescope with the email address used in Canvas).
    - Upload the exported zip file package to Gradescope

Congratulations on completing your first assignment!
</div>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 10)
```
>>> square_num(2)
4
```
failure_message: The function should return square of 2, which is 4.
success_message: First test (visible) is correct!


Tests for q2 (points: 10)
```
>>> fruit_by_color.index.to_list()
[1, 3, 0, 2]
```
